---
title: 'Latent Semantic Analysis'
description: 'Build upon your NLP feature engineering knowledge and implement LSA using SciKit-Learn.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Latent Semantic Analysis

```yaml
type: VideoExercise
key: 46e9015005
xp: 50
```

`@projector_key`
4ab1fcc06c5653e319cae3c2c5de82da
